---
title: 'Gold/Copper Ratio Analysis '
description: "The copper-gold ratio is calculated by dividing the market price of copper by the market price of gold. Copper is an industrial metal. It is used in plumbing, electric wiring, and its anti-bacterial properties coupled with its malleability lend the metal to be used in medical equipment. Demand increases during periods when economic output is rising. Gold is a store of value. It is molded into ingots, bullion, or displayed as jewelry. Its actual industrial applications are limited. The differing uses of the metals has allowed the copper-gold ratio to act as an accurate barometer of global growth.\n"
---

## An exercise title written in sentence case

```yaml
type: NormalExercise
key: c909bf2ad1
lang: python
xp: 100
skills: 2
```

This is the assignment text. It should help provide students with the background information needed.
The instructions that follow should be in bullet point form with clear guidance for what is expected.

`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`
```{python}
# Load datasets and packages here.
```

`@sample_code`
```{python}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{python}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{python}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: d6b89e460c
xp: 50
```

`@projector_key`
5468170f5429d871381e69a2ccc61227
